Global transcription factors analyses reveal hierarchy and synergism of regulatory networks and master virulence regulators in Pseudomonas aeruginosa

Author response:

The following is the authors’ response to the original reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

This work done by Huang et.al. revealed the complex regulatory functions and transcription network of 172 unknown transcription factors of Pseudomonas aeruginosa PAO1. The authors utilized ChIP-seq to profile TFs binding site information across the genome, demonstrating diverse regulatory relationships among them via hierarchical networks with three levels. They further constructed thirteen ternary regulatory motifs in small subs and co-association atlas with 7 core associated clusters. The study also uncovered 24 virulence-related master regulators. The pan-genome analysis uncovered both the conservation and evolution of TFs with P. aeruginosa complex and related species. Furthermore, they established a web-based database combining both existing and novel data from HT-SELEX and ChIP-seq to provide TF binding site information. This study offered valuable insights into studying transcription regulatory networks in P. aeruginosa and other microbes.

Strengths:

The results are presented with clarity, supported by well-organized figures and tables that not only illustrate the study's findings but also enhance the understanding of complex data patterns.

Thank you for your valuable feedback on our paper exploring the transcription regulatory networks in P. aeruginosa.

Weaknesses:

The results of this manuscript are mainly presented in systematic figures and tables. Some of the results need to be discussed as an illustration how readers can utilize these datasets.

We appreciate the valuable suggestion about enhancing the practical aspects of our manuscript. We have expanded the discussion section to include more detailed explanations of how these datasets can be utilized in practical applications.

Reviewer #2 (Public review):

In this work, the authors comprehensively describe the transcriptional regulatory network of Pseudomonas aeruginosa through the analysis of transcription factor binding characteristics. They reveal the hierarchical structure of the network through ChIP-seq, categorizing transcription factors into top-, middle-, and bottom-level, and reveal a diverse set of relationships among the transcription factors. Additionally, the authors conduct a pangenome analysis across the Pseudomonas aeruginosa species complex as well as other species to study the evolution of transcription factors. Moreover, the authors present a database with new and existing data to enable the storage and search of transcription factor binding sites. The findings of this study broaden our knowledge on the transcriptome of P. aeruginosa. This study sheds light on the complex interconnections between various cellular functions that contribute to the pathogenicity of P. aeruginosa, along with the associated regulatory mechanisms. Certain findings, such as the regulatory tendencies of DNA-binding domain-types, provides valuable insights on the possible functions of uncharacterized transcription factors and new functions of those that have already been characterized. The techniques used hold great potential for discovery of transcription factor functions in understudied organisms as well.

The study would benefit from a more clear discussion on the implications of various findings, such as binding preferences, regulatory preferences, and the link between regulatory crosstalk and virulence. Additionally, the pangenome analysis would be furthered through a discussion of the divergence of the transcription factors of P. aeruginosa PAO1 across species in relation to the findings on the hierarchical structure of the transcriptional regulatory network.

Thank you for your positive feedback and suggestions.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Major:

(1) It appears that many TFs are conserved among bacteria, archaebacteria, fungi, plants, and animals. Does this mean these TFs in bacterial could be the ancestors of TFs in fungi, plants, and animals? If we fetch these TFs out and build an evolutionary tree, can we visual the three kingdoms as well?

Thank you for this comment. While many TFs are conserved across bacteria, archaea, fungi, plants, and animals, this conservation does not necessarily imply a direct ancestral relationship. Instead, it may reflect the fundamental importance of certain domains and regulatory mechanisms, which could have arisen from a common ancestral system or through convergent evolution. If we fetch TF PA2032 out to build an evolutionary tree by setting PAO1 as the root, we can visualize these kingdoms in a tree. We added this content in the revised manuscript. Please see Figure S7D and Lines 404-411.

“The phylogenetic tree of PA2032 across bacteria, archaea, fungi, plants, and animals, with PAO1 as the root revealed that the bacterial TFs (purple) indicates a high degree of conservation within prokaryotes, suggesting a fundamental role in core regulatory processes. In contrast, eukaryotic TFs (fungi, plants, and animals) form distinct clades with longer branch lengths, indicating significant divergence and specialization during eukaryotic evolution. These findings suggest that while TF is conserved across domains of life, its functional roles and regulatory mechanisms have undergone substantial diversification in eukaryotes.”

(2) Can the authors give an indication how could we employ the findings of this study in designing next generation of antimicrobial agents?

Thank you for this important suggestion. We have provided this content in the discussion part. Please see Lines 481-492.

“The extensive datasets generated in this study offer valuable insights into understanding and targeting P. aeruginosa pathogenicity. The genome-wide binding profiles can be systematically analyzed through our hierarchical regulatory network framework to decode complex virulence mechanisms. The virulence-related master regulators and core regulatory clusters identified in this study highlighted key nodes of transcriptional control. Understanding these regulatory relationships is particularly valuable for identifying targets whose modulation would significantly impact virulence while accounting for potential compensatory mechanisms. This knowledge base thus provides a foundation for developing targeted approaches to combat P. aeruginosa infections, moving beyond traditional antibiotic strategies toward more sophisticated interventions based on regulatory network manipulation.”

Minor：

(1) Lines 178-180: It would strengthen the discussion to include a few additional references that support the claims made in this section, providing a more comprehensive context for the readers.

Yes. We have added more citations(1-5) (No. 1-5 in the references at the end of the rebuttal) to support the claims. Please see Line 182.

(2) Line 198: You mention 'seven' motifs containing toggle switches, but Fig.3 actually displays eight motifs. Please revise this discrepancy to ensure consistency between the text and the figure.

Yes. We have revised the wording to “eight”. Please see Line 200.

(3) Figure 3A: Consider adding a diagram or legend that represents the colors associated with each DNA-binding domain (DBD) family.

Thank you for your suggestion. The colors of DBD were aligned with the legend in Figure S3. We have added it in Figure 3A.

Reviewer #2 (Recommendations for the authors):

Line 21: The use of the abbreviation 'TF' should be done at the first instance of 'transcription factor'.

Yes. We have revised it. Please see Line 21.

Line 74: The purpose of this paragraph is slightly unclear. It is recommended that appropriate modifications are made.

We are sorry for the confusion. The purpose of this paragraph was to introduce the major virulence pathways in P. aeruginosa and mention the important role of TRN in these pathways. We have modified it to make it clearer. Please see Lines 74-75.

“P. aeruginosa employs diverse virulence pathways to establish successful infection, with QS being one of the major mechanisms involving the expression of many virulence genes.”

Line 113: How were these 172 TFs selected?

Thank you for indicating this question. In a previous study, we performed HT-SELEX to characterize the DNA-binding motifs of all TFs in P. aeruginosa PAO1, successfully identifying binding sequences for 182 TFs. To further elucidate the binding landscapes of the rest, we performed ChIP-seq on the remaining TFs (172 TFs in total with high-quality ChIP-seq libraries). Please see Lines 100-101 in the revised manuscript.

Line 119: Defining other features, namely downstream and include Feature, would be helpful.

Thank you for your suggestion. We have added the definition for all peak annotation in the legend. Please see Lines 569-574.

“Annotation heatmap of all peak distribution with 6 locations: Upstream, where the peak is located entirely upstream of the gene; Downstream, where the peak is positioned completely downstream of the gene; Inside, where the peak is entirely contained within the gene body; OverlapStart, where the peak overlaps with the 5' end of the gene; OverlapEnd, where the peak overlaps with the 3' end of the gene; and IncludeFeature, where the peak completely encompasses the gene.”

Line 129: The distribution type of AraC-type TFs is unclear - it is mentioned that AraC has a 'broad distribution', but it is later stated that it has a 'narrow distribution'.

We are sorry for this mistake, and we have revised the example for “broad distribution”, which is Cor_CI instead of AraC. Please see Lines 132-135.

Line 161: 'h value' here may need to be modified to 'absolute h value'.

Yes. We have revised it. Please see Line 164.

Line 502: "s The DNA" needs to be corrected.

Yes. We have revised it. Please see Line 514.

Line 515: It would be helpful to readers if the reference used for these pathways was cited.

Yes. We have added the review reference (Shao et al, 2023) related to these pathways(6) (the 6th reference at the end of the rebuttal). Please see Line 527.

Line 558: "Translation start site" needs to be corrected to "Transcription start site"

The “TSS” here exactly indicated “Translation start site”.

Line 593. "Virulent" pathways needs to be corrected to "virulence" pathways.

Yes. We have revised it. Please see Line 609.

Line 604: The type of categorization based on which the proportion of genes is displayed needs to be mentioned.

Yes, we agree. We have added the type of categorization in the legend. Please see Lines 621-627.

“Figure 6. Conservation and variability of TFs in PAO1. (A). The pie chart shows the proportions of genes categorized by their presence across P. aeruginosa strains for all genes. (B). The pie chart shows the distribution of TFs identified from PAO1 across different conservation categories. (C). The bar plot of the proportion for non-core TFs. Genes are categorized based on their presence frequency across P. aeruginosa strains: Core genes (present in 99% ~ 100% strains), Soft core genes (present in 95% ~ 99% strains), Shell genes (present in 15% ~ 95% strains), and Cloud genes (present in 0% ~ 15% strains).”

Reference:

(1) Liang H, Deng X, Li X, Ye Y, Wu M. 2014. Molecular mechanisms of master regulator VqsM mediating quorum-sensing and antibiotic resistance in Pseudomonas aeruginosa. Nucleic acids research 42:10307-10320.

(2) Jones CJ, Ryder CR, Mann EE, Wozniak DJ. 2013. AmrZ modulates Pseudomonas aeruginosa biofilm architecture by directly repressing transcription of the psl operon. Journal of bacteriology 195:1637-1644.

(3) Hickman JW, Harwood CS. 2008. Identification of FleQ from Pseudomonas aeruginosa as ac‐di‐GMP‐responsive transcription factor. Molecular microbiology 69:376-389.

(4) Déziel E, Gopalan S, Tampakaki AP, Lépine F, Padfield KE, Saucier M, Xiao G, Rahme LG. 2005. The contribution of MvfR to Pseudomonas aeruginosa pathogenesis and quorum sensing circuitry regulation: multiple quorum sensing‐regulated genes are modulated without affecting lasRI, rhlRI or the production of N‐acyl‐L‐homoserine lactones. Molecular microbiology 55:998-1014.

(5) Lizewski SE, Lundberg DS, Schurr MJ. 2002. The transcriptional regulator AlgR is essential for Pseudomonas aeruginosa pathogenesis. Infection and immunity 70:6083-6093.

(6) Shao X, Yao C, Ding Y, Hu H, Qian G, He M, Deng X. 2023. The transcriptional regulators of virulence for Pseudomonas aeruginosa: Therapeutic opportunity and preventive potential of its clinical infections. Genes & Diseases 10:2049-2063.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The transcription factor regulatory network in Pseudomonas aeruginosa is complex and involves multiple regulators that respond to various environmental signals and physiological cues by regulating gene expression. However, the biological functions of at least half of its 373 putative transcription factors (TFs) remain uncharacterised. Herein, chromatin immunoprecipitation sequencing (ChIP-seq) was used to investigate the binding sites of 172 TFs in the P. aeruginosa PAO1 strain. The results revealed 81,009 significant binding peaks in the genome, more than half of which were located in the promoter regions. To further decode the diverse regulatory relationships among TFs, a hierarchical network was assembled into three levels: top, middle, and bottom. Thirteen ternary regulatory motifs revealed flexible relationships among TFs in small hubs, and a comprehensive co-association atlas was established, showing the enrichment of seven core associated clusters. Twenty-four TFs were identified as the master regulators of virulence-related pathways. The pan-genome analysis revealed the conservation and evolution of TFs in P. aeruginosa complex and other species. A Web-based database combining existing and new data from ChIP-seq and the high-throughput systematic evolution of ligands by exponential enrichment was established for searching TF-binding sites. This study provides important insights into the pathogenic mechanisms of P. aeruginosa and related bacteria and is expected to contribute to the development of effective therapies for infectious diseases caused by this pathogen.

Revised version: The transcription factor (TF) regulatory network in Pseudomonas aeruginosa is complex and involves multiple regulators that respond to various environmental signals and physiological cues by regulating gene expression. However, the biological functions of at least half of its 373 putative TFs remain uncharacterised. Herein, chromatin immunoprecipitation sequencing (ChIP-seq) was used to investigate the binding sites of 172 TFs in the P. aeruginosa PAO1 strain. The results revealed 81,009 significant binding peaks in the genome, more than half of which were located in the promoter regions. To further decode the diverse regulatory relationships among TFs, a hierarchical network was assembled into three levels: top, middle, and bottom. Thirteen ternary regulatory motifs revealed flexible relationships among TFs in small hubs, and a comprehensive co-association atlas was established, showing the enrichment of seven core associated clusters. Twenty-four TFs were identified as the master regulators of virulence-related pathways. The pan-genome analysis revealed the conservation and evolution of TFs in P. aeruginosa complex and other species. A Web-based database combining existing and new data from ChIP-seq and the high- throughput systematic evolution of ligands by exponential enrichment was established for searching TF-binding sites. This study provides important insights into the pathogenic mechanisms of P. aeruginosa and related bacteria and is expected to contribute to the development of effective therapies for infectious diseases caused by this pathogen.

Introduction

Transcription factors (TFs) are key regulators of transcription and thus play crucial roles in mediating multiple biological pathways and events in eukaryotes and prokaryotes. Based on their DNA-binding domains (DBDs), TFs can be classified into different families, such as the LysR, AraC, and LuxR families1. TFs can activate and repress the expression of a set of genes in response to various environmental and/or specific signal triggers2,3. Several approaches can be used to study TFs, including chromatin immunoprecipitation sequencing (ChIP-seq)4 and ChIP-exo5 in in vivo research and DNA affinity purification sequencing (DAP-seq)6 and high-throughput systematic evolution of ligands by exponential enrichment (HT-SELEX)7 in in vitro research. ChIP-seq is a powerful in vivo technique that enables more accurate genome-wide mapping of TF–DNA interactions than in vitro methods because TFs may interact with other co-regulators in an environment-specific pattern, thereby altering binding preferences in vivo8,9.

These sequencing approaches have been used to illustrate the biological functions of TFs in several bacterial species, including Mycobacterium tuberculosis, Vibrio cholerae, Salmonella enterica, Escherichia coli, and Clostridium thermocellum10–15. In addition, several databases and models have been constructed based on experimental validation or machine learning for studying TFs in bacteria, such as CollecTF, PredicTF, PRODORIC, DBTBS, and RegulomePA16–20. However, the functions and spatiotemporal interactions of most bacterial TFs remain to be clarified.

The gram-negative bacterium Pseudomonas aeruginosa is a major human opportunistic pathogen that grows ubiquitously. P. aeruginosa tends to cause infection in burn victims, patients with cystic fibrosis (CF), and individuals hospitalised for long periods21,22. It utilises versatile pathways to exert its virulence, such as biofilm formation, quorum sensing (QS), type III and type VI secretion systems (T3SS and T6SS, respectively), motility, siderophore production, oxidative stress resistance, and antibiotic resistance, all of which are under the control of a complicated TF regulatory network (TRN). The metabolic changes regulated by the TRN enhance the pathogen’s adaptability under different infection and survival conditions.

Earlier version: QS, a major virulence pathway in P. aeruginosa, is a bacterial cell–cell communication mechanism involving the expression of many virulence genes. Its regulation occurs hierarchically, comprising the interconnected las, rhl, and pqs systems23,24. Through its transcriptional regulator, LasR, and in collaboration with its corresponding autoinducer, 3-oxo-C12-HSL (encoded by lasI), the las system triggers the activation of the rhl and pqs systems23,25. In addition, P. aeruginosa attaches to and colonises surfaces primarily using two forms of motility: (i) swarming and swimming powered by flagella, which are multicellular and individual cell movements, respectively, and (ii) twitching powered by type 4 pili (T4P)26, wherein T4P promote surface attachment by twitching to pull the cell closer to the attachment sites27. The motility of P. aeruginosa is thus the initial step in its pathogenic process and is associated with the upregulation of virulence and the induction of host defence. Overall, the virulence-related pathways, including QS and motility, are controlled by the sophisticated TRN in P. aeruginosa.

Revised version: P. aeruginosa employs diverse virulence pathways to establish successful infection, with QS being one of the major mechanisms involving the expression of many virulence genes. Its regulation occurs hierarchically, comprising the interconnected las, rhl, and pqs systems23,24. Through its transcriptional regulator, LasR, and in collaboration with its corresponding autoinducer, 3-oxo-C12-HSL (encoded by lasI), the las system triggers the activation of the rhl and pqs systems23,25. Additionally, P. aeruginosa utilizes motility systems for surface attachment and colonization, including: (i) swarming and swimming powered by flagella, which are multicellular and individual cell movements, respectively, and (ii) twitching powered by type 4 pili (T4P)26, wherein T4P promote surface attachment by twitching to pull the cell closer to the attachment sites27. The motility of P. aeruginosa is thus the initial step in its pathogenic process and is associated with the upregulation of virulence and the induction of host defence. Both QS and motility pathways are precisely controlled by the sophisticated TRN in P. aeruginosa. Understanding how these virulence pathways are integrated and regulated is crucial for developing novel therapeutic strategies against P. aeruginosa infections.

Earlier version: Despite the TRN being essential for understanding how bacterial genes are regulated synergistically, only a limited number of binding targets of TFs have been identified in P. aeruginosa over the past decades. For example, ChIP-seq and RNA sequencing (RNA-seq) co-analysis has been used to map a P. aeruginosa genomic regulatory network (PAGnet) of 20 key virulence-related TFs, while HT-SELEX has been used to characterise the DNA-binding specificities of 182 TFs28. DAP-seq has been used to reveal a regulatory network of 55 response regulators (RRs) of two-component systems in P. aeruginosa, 51 of which are from the PAO1 strain29. However, almost half of the TFs remain to be annotated with biological functions, and their downstream targets and upstream regulators remain unknown. To fill this knowledge gap, we performed ChIP-seq of 172 TFs (Table S1) from P. aeruginosa, which revealed a hierarchical and co-association regulatory network among TFs and newly identified virulence-related regulators. We also established a Web-based database to store and search TF-binding patterns together with ChIP-seq and HT-SELEX data, which is available to the public. Our findings and tools will significantly contribute to future mechanistic studies on and a comprehensive understanding of the TRN in P. aeruginosa and other bacterial species.

Revised version: Despite the TRN being essential for understanding how bacterial genes are regulated synergistically, only a limited number of binding targets of TFs have been identified in P. aeruginosa over the past decades. For example, ChIP-seq and RNA sequencing (RNA-seq) co-analysis has been used to map a P. aeruginosa genomic regulatory network (PAGnet) of 20 key virulence-related TFs, while HT-SELEX has been used to characterise the DNA-binding specificities of 182 TFs28. DAP-seq has been used to reveal a regulatory network of 55 response regulators (RRs) of two-component systems in P. aeruginosa, 51 of which are from the PAO1 strain29. However, almost half of the TFs remain to be annotated with biological functions, and their downstream targets and upstream regulators remain unknown. To fill this knowledge gap, we performed ChIP-seq of 172 TFs, which are the most rest of the TFs after HT- SELEX screening (Table S1) from P. aeruginosa, which revealed a hierarchical and co-association regulatory network among TFs and newly identified virulence-related regulators. We also established a Web-based database to store and search TF-binding patterns together with ChIP-seq and HT-SELEX data, which is available to the public. Our findings and tools will significantly contribute to future mechanistic studies on and a comprehensive understanding of the TRN in P. aeruginosa and other bacterial species.

Results

ChIP-seq of 172 uncharacterised TFs reveals a global transcriptional atlas of P. aeruginosa

In this study, 373 proteins in P. aeruginosa were annotated as TFs with DNA-binding activity based on the existing annotations in the Pseudomonas Genome Database30 (Fig. 1A). To profile the whole TF–DNA binding landscape and construct a comprehensive regulatory network of P. aeruginosa, ChIP-seq experiments were performed on 172 TFs to determine the TF-binding sites (TFBSs) (Fig. 1A). Libraries were constructed using vesicular stomatitis virus glycoprotein tagging, and raw sequence reads were mapped to the PAO1 genome using bowtie2 (Version 2.3.4.1)31. Using MACS232, 81,009 significant binding peaks with a cut-off p-value of 0.001 were identified. Subsequently, the R package ChIPpeakAnno33 was used to define the peak locations and find the nearest genes. The peak location can be divided into six features: upstream, overlapStart, inside, overlapEnd, downstream, and includeFeature. The percentages of binding peak locations for each TF indicated that more than half of the TFs preferentially bind to the upstream (upstream intergenic regions) and overlapStart (overlapped with the transcription start sites [TSSs] of genes) regions (Fig. 1B). Apart from these regions, some TFs bind to the inside regions (coding regions of genes). Few TFs were found to preferentially bind to the rest of the peak location features.

Overview of ChIP-seq results.

Earlier version: (A). Density of all TFs (green) and ChIPed TFs (orange) in this study throughout the P. aeruginosa genome. (B). Annotation heatmap of all peak distribution with 6 locations: Upstream, OverlapStart, Inside, OverlapEnd, Downstream, and IncludeFeature. (C). Peak distance to the translational start site (TSS) of each DBD family. (D). Treemap of the 172 TFs peak numbers based on DBD family. Each box’s size represents the family’s size (number of peaks), and the explained variance of each DBD type means the color shades of each box. DBD families of ChIPed TFs are classified into 20 different categories: LuxR, LysR, Two DBDs (Two DNA binding domains), AraC, TetR, ArsR, CRP, OmpR, GntR, MarR, AsnC, Cro/CI, TyrR, Rrf2, MerR, IclR, Fis, RpiR, DeoR, and undetectable. (E). The dot plot shows the top 10 gene ontology (GO) terms from the PseudoCAP annotation of ChIPed TFs. The size of the dots indicates the significance of each functional category, quantified by −log10 (p. adjust).

Revised version: (A). Density of all TFs (green) and ChIPed TFs (orange) in this study throughout the P. aeruginosa genome. (B). Annotation heatmap of all peak distribution with 6 locations: Upstream, where the peak is located entirely upstream of the gene; Downstream, where the peak is positioned completely downstream of the gene; Inside, where the peak is entirely contained within the gene body; OverlapStart, where the peak overlaps with the 5’ end of the gene; OverlapEnd, where the peak overlaps with the 3’ end of the gene; and IncludeFeature, where the peak completely encompasses the gene. (C). Peak distance to the translational start site (TSS) of each DBD family. (D). Treemap of the 172 TFs peak numbers based on DBD family. Each box’s size represents the family’s size (number of peaks), and the explained variance of each DBD type means the color shades of each box. DBD families of ChIPed TFs are classified into 20 different categories: LuxR, LysR, Two DBDs (Two DNA binding domains), AraC, TetR, ArsR, CRP, OmpR, GntR, MarR, AsnC, Cro/CI, TyrR, Rrf2, MerR, IclR, Fis, RpiR, DeoR, and undetectable. (E). The dot plot shows the top 10 gene ontology (GO) terms from the PseudoCAP annotation of ChIPed TFs. The size of the dots indicates the significance of each functional category, quantified by −log10 (p. adjust).

Earlier version: To further provide insights into the functional roles of different TFs in gene regulation, we analysed the peak distance to the TSS, which revealed several patterns (Fig. 1C). First, most of the TFs exhibited peak distributions that were concentrated in the proximal region around the TSS. Some TFs, such as ArsR and AraC, showed a broader distribution of peak distance, indicating that these TFs might also be active at positions farther from the TSS. Furthermore, certain TFs, including CRP and DeoR, displayed distinct bimodal or multimodal distributions, while other TFs, such as Fis and AraC, showed smoother and more narrowly focused distributions, suggesting their different binding preferences. The number of peaks of each TF and its DBD types are summarised in a treemap in Fig. 1D. TFs in the LysR and AraC families accounted for almost half of all of the peaks. PA2718, classified in the MerR family, yielded the highest number (2,480) of binding peaks, followed by the TF PA0756 (2,080 binding peaks) from the OmpR family. Furthermore, eight TFs showed no more than 10 peaks, and the TFs PA4074 and PA4806 had only two peaks each. To preliminarily identify the different regulatory functions of TFs via their different binding preferences, we performed Gene Ontology (GO) functional enrichment analysis for each TF and presented the top 10 GO terms (Fig. 1E). The significantly enriched GO terms (BH-adjusted P < 0.05) for all of the 172 TFs analysed by ChIP-seq are listed in Table S2, and these terms suggest divergent functions of the TFs. Of these functional categories, most were related to metabolism, such as transcription, translation, ribosome structure, and GTPase activity, while some of the functional categories were associated with virulence, including T4P, O antigen biosynthesis, and biofilm formation (Fig. 1E).

Revised version: To further provide insights into the functional roles of different TFs in gene regulation, we analysed the peak distance to the TSS, which revealed several patterns (Fig. 1C). First, most of the TFs exhibited peak distributions that were concentrated in the proximal region around the TSS. Some TFs, such as Cor_CI, showed a broader distribution of peak distance, indicating that these TFs might also be active at positions farther from the TSS. Furthermore, certain TFs, including RpiR and DeoR, displayed distinct bimodal or multimodal distributions, while other TFs, such as LysR and AraC, showed smoother and more narrowly focused distributions, suggesting their different binding preferences. The number of peaks of each TF and its DBD types are summarised in a treemap in Fig. 1D. TFs in the LysR and AraC families accounted for almost half of all of the peaks. PA2718, classified in the MerR family, yielded the highest number (2,480) of binding peaks, followed by the TF PA0756 (2,080 binding peaks) from the OmpR family. Furthermore, eight TFs showed no more than 10 peaks, and the TFs PA4074 and PA4806 had only two peaks each. To preliminarily identify the different regulatory functions of TFs via their different binding preferences, we performed Gene Ontology (GO) functional enrichment analysis for each TF and presented the top 10 GO terms (Fig. 1E). The significantly enriched GO terms (BH-adjusted P < 0.05) for all of the 172 TFs analysed by ChIP- seq are listed in Table S2, and these terms suggest divergent functions of the TFs. Of these functional categories, most were related to metabolism, such as transcription, translation, ribosome structure, and GTPase activity, while some of the functional categories were associated with virulence, including T4P, O antigen biosynthesis, and biofilm formation (Fig. 1E).

Hierarchical networks of TFs based on pairwise interactions

TF-coding genes can also serve as targets for other TFs, leading to interactions between the DNA-binding profiles of TFs. To clarify these relationships, lists of binding targets of the TFs obtained using ChIP-seq and HT-SELEX were used to assemble a system-level hierarchical regulatory network. After filtering by TFs that bind to promoter regions, 947 unique pairwise patterns between TFs were identified from a total of 13,375 promoter interactions.

Earlier version: We next defined a statistic index, h, based on the out-degree (O) and in-degree (I) of interactions to measure the direction and hierarchical level of TFs [h = (O-I)/(O+I)] in our established hierarchical regulatory network and thus identify hubs and information-flow bottlenecks34. The h value of each TF can reflect the direction and extent of information flow. A positive h indicates that the TF acts ‘upstream’, whereas a negative value suggests that it acts ‘downstream’. The higher the h value, the stronger the implication of the direction. A density plot of all of the TFs with their h values revealed that they can be approximately divided into three groups: top (h ∈ [0.75, 1]), middle (h ∈ [-0.75, 0.75]), and bottom (h ∈ [-1, 0.75]) levels (Fig. S2A). The top-level TFs are those that tend to control many other TFs, while the bottom-level TFs are those that are more regulated by other TFs than acting as regulators themselves. The middle-level TFs are those that connect the top- and bottom-level TFs. Of the 231 TFs in this hierarchy, 100, 46, and 85 TFs were defined as top-, middle-, and bottom-level TFs, respectively. The complete hierarchical information of all of the TFs is provided in Table S3. To avoid extreme cases and aid visualisation, we summarised the hierarchical regulatory network by removing TFs for which (O+I) was more than 10 (Fig. 2A and Fig. S2B). The network was visualised using Cytoscape software35.

Revised version: We next defined a statistic index, h, based on the out-degree (O) and in-degree (I) of interactions to measure the direction and hierarchical level of TFs [h = (O-I)/(O+I)] in our established hierarchical regulatory network and thus identify hubs and information-flow bottlenecks34. The h value of each TF can reflect the direction and extent of information flow. A positive h indicates that the TF acts ‘upstream’, whereas a negative value suggests that it acts ‘downstream’. The higher the absolute h value, the stronger the implication of the direction. A density plot of all of the TFs with their h values revealed that they can be approximately divided into three groups: top (h [0.75, 1]), middle (h [-0.75, 0.75]), and bottom (h [-1, 0.75]) levels (Fig. S2A). The top-level TFs are those that tend to control many other TFs, while the bottom- level TFs are those that are more regulated by other TFs than acting as regulators themselves. The middle-level TFs are those that connect the top- and bottom-level TFs. Of the 231 TFs in this hierarchy, 100, 46, and 85 TFs were defined as top-, middle-, and bottom-level TFs, respectively. The complete hierarchical information of all of the TFs is provided in Table S3. To avoid extreme cases and aid visualisation, we summarised the hierarchical regulatory network by removing TFs for which (O+I) was more than 10 (Fig. 2A and Fig. S2B). The network was visualised using Cytoscape software35.

Hierarchical networks and network motifs.

(A). Overview of hierarchy regulatory network after removing TFs with degree O+I <= 10. Nodes indicate TFs, and the color represents the hierarchical level. The top level was highlighted in red, the middle in yellow, and the bottom highlighted in blue. The edges with arrows indicate the regulatory direction. Gray means downward-pointing, and red means upward-pointing. (B). The auto-regulator motif of 9 TFs. (C). Three transcription- factor motifs occurrences with five basic triangular motifs and eight toggle switch motifs. The circle presents TF, and the arrow indicates the regulatory direction. (D). Alluvial diagram reveals basic triangle motif 3 (n = 1956) depending on DBD types. The color of splines is highlighted in different DBD families, and the name of DBD families are labeled.

Earlier version: Additionally, we investigated the factors influencing the hierarchical level of TFs, such as DBD types (e.g., DBD families containing more TFs). TFs from the AraC, GntR, and TetR families tended to cluster at the top level. TFs from the MarR family and two-DBD-type TFs (i.e., TFs with two DBDs) were mostly found in the bottom level. Additionally, among the characterised TFs, 58% were clustered at the bottom level, probably because these TFs directly regulated the downstream genes involved in important phenotypes, such as AlgR, AmrZ, MvfR, FleQ, and VqsM. Overall, the organised hierarchical regulatory network profiling at the three levels showed complex pairwise interactions among TFs in P. aeruginosa.

Revised version: Additionally, we investigated the factors influencing the hierarchical level of TFs, such as DBD types (e.g., DBD families containing more TFs). TFs from the AraC, GntR, and TetR families tended to cluster at the top level. TFs from the MarR family and two-DBD-type TFs (i.e., TFs with two DBDs) were mostly found in the bottom level. Additionally, among the characterised TFs, 58% were clustered at the bottom level, probably because these TFs directly regulated the downstream genes involved in important phenotypes, such as AlgR36, AmrZ37, MvfR38, FleQ39, and VqsM40. Overall, the organised hierarchical regulatory network profiling at the three levels showed complex pairwise interactions among TFs in P. aeruginosa.

Ternary regulatory motifs show flexible relationships among TFs in P. aeruginosa

Earlier version: Apart from their global hierarchical regulatory structure, we investigated the constituent regulatory network motifs of TFs, which revealed small connectivity patterns associated with canonical functions36. Positive and negative feedback loops are prevalent in bacterial regulatory networks, which generally involve combinations of several genes to appropriately respond to stimuli. One of the regulatory patterns was a single auto-regulator that can self-regulate its expression or activity. Through analysis, nine auto-regulator TFs were identified (Fig. 2B). Compared with non-auto- regulators, auto-regulators usually tend to be repressors, which play important roles in maintaining a steady state37,38. For example, the key T3SS activators PsrA and HrpL can repress their own expression in P. aeruginosa and P. savastanoi, respectively39–41.

Revised version: Apart from their global hierarchical regulatory structure, we investigated the constituent regulatory network motifs of TFs, which revealed small connectivity patterns associated with canonical functions41. Positive and negative feedback loops are prevalent in bacterial regulatory networks, which generally involve combinations of several genes to appropriately respond to stimuli. One of the regulatory patterns was a single auto-regulator that can self-regulate its expression or activity. Through analysis, nine auto-regulator TFs were identified (Fig. 2B). Compared with non-auto- regulators, auto-regulators usually tend to be repressors, which play important roles in maintaining a steady state42,43. For example, the key T3SS activators PsrA and HrpL can repress their own expression in P. aeruginosa and P. savastanoi, respectively44–46.

Earlier version: In addition, the most fundamental regulatory pattern was the ternary TF motif. We computed all possible motifs, which are listed in Fig. 2C and Table S4. The basic triangular motifs were monodirectional regulatory structures with five different motifs. The most frequently observed motif (6,535 occurrences) was motif_2, wherein two different TFs co-regulate another TF. The remaining seven types of motifs, which contained toggle switches, occurred much less frequently than the basic motifs. The more complex the relationships among TFs, the lesser their occurrence. In particular, motif_10, motif_12, and motif_13, characterised by two or three mutually regulating TFs, were not found even once (Fig. 2C). The high occurrence of different TFs in diverse motifs revealed a complex and flexible regulatory network in P. aeruginosa. Next, we constructed a basic motif_3 regulatory network based on potential TF–TF interactions to detect the regulatory preferences based on DBD types (Fig. 2D). Basic motif_3 was a typical hierarchical regulatory network in which TF1 regulates TF2, which in turn regulates TF3. We found that the AraC, Fis, and GntR families tended to harbour top-level TFs rather than middle- and bottom-level TFs, while the OmpR family tended to harbour TFs from middle- and bottom-levels more than those from the top level. Furthermore, the LysR and LuxR families revealed non-significant differences among these three TF levels (Fig. 2D). Taken together, our integrated hierarchical regulatory profiling revealed a multifaceted TF–TF connection, reflecting the regulatory preferences of TFs and DBD types in P. aeruginosa.

Revised version: In addition, the most fundamental regulatory pattern was the ternary TF motif. We computed all possible motifs, which are listed in Fig. 2C and Table S4. The basic triangular motifs were monodirectional regulatory structures with five different motifs. The most frequently observed motif (6,535 occurrences) was motif_2, wherein two different TFs co-regulate another TF. The remaining eight types of motifs, which contained toggle switches, occurred much less frequently than the basic motifs. The more complex the relationships among TFs, the lesser their occurrence. In particular, motif_10, motif_12, and motif_13, characterised by two or three mutually regulating TFs, were not found even once (Fig. 2C). The high occurrence of different TFs in diverse motifs revealed a complex and flexible regulatory network in P. aeruginosa. Next, we constructed a basic motif_3 regulatory network based on potential TF–TF interactions to detect the regulatory preferences based on DBD types (Fig. 2D). Basic motif_3 was a typical hierarchical regulatory network in which TF1 regulates TF2, which in turn regulates TF3. We found that the AraC, Fis, and GntR families tended to harbour top-level TFs rather than middle- and bottom-level TFs, while the OmpR family tended to harbour TFs from middle- and bottom-levels more than those from the top level. Furthermore, the LysR and LuxR families revealed non-significant differences among these three TF levels (Fig. 2D). Taken together, our integrated hierarchical regulatory profiling revealed a multifaceted TF–TF connection, reflecting the regulatory preferences of TFs and DBD types in P. aeruginosa.

Clustering of 103 TF-binding motifs

Earlier version: According to their binding sites, a total of 103 binding motifs were determined. The TFs were clustered based on the positional weight matrix similarity of their motifs. A circular phylogenetic tree was constructed using hierarchical clustering of the pairwise similarity matrix (Fig. S3 and Table S5). Each node in the tree represents a TF, colour-highlighted according to its DBD family classification. The clustering analysis revealed four distinct clusters of TFs, one each coloured light green (Cluster 1), light blue (Cluster 2), purple (Cluster 3), and pink (Cluster 4). These clusters indicate groups of TFs with similar binding properties or regulatory functions.

Revised version: According to their binding sites, a total of 103 binding motifs were determined. The TFs were clustered based on the positional weight matrix similarity of their motifs. A circular phylogenetic tree was constructed using hierarchical clustering of the pairwise similarity matrix (Fig. S3 and Table S5). Each node in the tree represents a TF, colour-highlighted according to its DBD family classification. The clustering analysis revealed four distinct clusters of TFs, one each coloured light green (Cluster 1), light blue (Cluster 2), purple (Cluster 3), and pink (Cluster 4). These clustersindicate groups of TFs with similar binding properties or regulatory functions.

Within each identified cluster, TFs tend to share conserved motif patterns, indicating potential functional relationships. For instance, it was found that the TFs in cluster 1 were enriched in adenine (A) and thymine (T) bases. Further, cluster 4 was more likely to have an inverted repeat (IR) sequence, such as PA0367 and PA1015. This conservation suggests that TFs within the same cluster co-regulate similar sets of genes. Additionally, the clustering uncovered novel associations between TFs from different DBD families, indicating potential new regulatory interactions or evolutionary links that warrant further investigation.

Genome-wide synergistic co-association of TFs in P. aeruginosa

Earlier version: To further investigate the crosstalk among TFs in P. aeruginosa, we computed the co- association scores based on the overlaps between the peaks of all pairs of TFs and ChIP-seq data (Fig. S4A and Table S6). Among a total of 16,175 TF pairs, we found 5,716 significant TF interactions with a co-association score of up to 0.1 based on an elbow statistic (Fig. S4B). The relationships among TFs can be further determined as core clusters if their co-association scores are more than 0.4. The clustering analysis of core networks using Glay clustering42 resulted in seven associated modules, each containing a diverse number of TFs with the DBD families highlighted in different colours (Fig. 3A).

Revised version: To further investigate the crosstalk among TFs in P. aeruginosa, we computed the co- association scores based on the overlaps between the peaks of all pairs of TFs and ChIP-seq data (Fig. S4A and Table S6). Among a total of 16,175 TF pairs, we found 5,716 significant TF interactions with a co-association score of up to 0.1 based on an elbow statistic (Fig. S4B). The relationships among TFs can be further determined as core clusters if their co-association scores are more than 0.4. The clustering analysis of core networks using Glay clustering47 resulted in seven associated modules, each containing a diverse number of TFs with the DBD families highlighted in different colours (Fig. 3A).

Core co-association regulatory networks.

(A). Core clusters of significant co-binding patterns of TFs. Each TF is highlighted in different colors based on DBD types. The co-association score by pair of TFs was calculated by Jaccard statistics, which measures the ratio of the number of base pairs in overlapped binding peaks on both TFs to the number of base pairs in their union. (B). The histogram’s overlapped target genes of TFs in cluster 3 represent the number of target peaks in the individual/overlapped set. (C). The network of co-regulation of TFs in cluster 3 with co-bound targets of more than 4. (D-E). Genome browser view of TFs in cluster 3 binding intensities at the PA2504 and pqsH locus.

Earlier version: Next, we summarised the potential functional crosstalk among the middle-scale co- associated five TFs (PA2417, PA2718, PA0756, PmiR, and AgtR) in cluster 3. Among a total of 1,241 target genes, 372 were co-bound by at least three TFs, and 35 were co- bound by all five TFs, indicating that these TFs have a high level of collaboration in regulatory patterns (Fig. 3B). There were smaller intersections among different combinations of TFs, highlighting potential co-regulation or shared binding sites. Binding targets co-bound by more than four TFs are summarised in Fig. 3C. The dense network and connectivity indicate a complex regulatory landscape, where these TFs may orchestrate a coordinated regulation of gene expression. For example, peak visualization of PA2504 revealed distinct binding patterns of these five TFs compared with the control group (Fig. 3D). PA2504 is the sole partner of the PppH RNA hydrolase related to several important cellular factors43. In addition, four of the five TFs (barring PA2718) can bind to the promoter region of pqsH, which is involved in the pqs pathway (Fig. 3E). Notably, PmiR has previously been verified to bind to the pqsH promoter region and thereby regulate virulence44. PA2504 and pqsH are highlighted in blue and green, respectively, in Fig. 3C.

Revised version: Next, we summarised the potential functional crosstalk among the middle-scale co- associated five TFs (PA2417, PA2718, PA0756, PmiR, and AgtR) in cluster 3. Among a total of 1,241 target genes, 372 were co-bound by at least three TFs, and 35 were co- bound by all five TFs, indicating that these TFs have a high level of collaboration in regulatory patterns (Fig. 3B). There were smaller intersections among different combinations of TFs, highlighting potential co-regulation or shared binding sites. Binding targets co-bound by more than four TFs are summarised in Fig. 3C. The dense network and connectivity indicate a complex regulatory landscape, where these TFs may orchestrate a coordinated regulation of gene expression. For example, peak visualization of PA2504 revealed distinct binding patterns of these five TFs compared with the control group (Fig. 3D). PA2504 is the sole partner of the PppH RNA hydrolase related to several important cellular factors48. In addition, four of the five TFs (barring PA2718) can bind to the promoter region of pqsH, which is involved in the pqs pathway (Fig. 3E). Notably, PmiR has previously been verified to bind to the pqsH promoter region and thereby regulate virulence49. PA2504 and pqsH are highlighted in blue and green, respectively, in Fig. 3C.

Twenty-four master regulators are associated with virulence

After elucidating the binding targets of 172 TFs, we investigated their potential functions in the virulence-related pathways and pathogenesis of P. aeruginosa. For this purpose, we performed an enrichment analysis using a hypergeometric test to define the master regulators (BH-adjusted P < 0.05) that might play important roles in regulating specific pathways. We primarily focused on TF-binding profiles involving the promoter regions of genes involved in nine pathways, namely QS, motility, biofilm production, antibiotic resistance, T6SS, T3SS, reactive oxygen species (ROS) resistance, pyocyanin production, and siderophore production. Accordingly, 134 of the 172 TFs were found to bind to at least one target gene associated with the abovementioned virulent pathways, and 24 were identified as master regulators of genes involved in six pathways, namely siderophore production, biofilm production, pyocyanin production, QS, motility, and ROS resistance pathways (Fig. 4A and Table S7).

Newly identified virulence-related master regulators.

Earlier version: (A). Overview of all identified master regulators related to 6 pathways, including QS, motility, biofilm, siderophore, pyocyanin, and ROS. Each circle represents one TF, with the height indicating the significance, quantified by -log10 (p.adjust), and the size indicates the number of targets associated with the virulence pathway. (B). Intersection of master regulators in 6 virulent pathways. The bar chart on the top shows the number of intersections, while the matrix below indicates which TFs are involved in specific biological processes such as siderophore production, pyocyanin production, biofilm formation, ROS response, QS, and motility.

Revised version: (A). Overview of all identified master regulators related to 6 pathways, including QS, motility, biofilm, siderophore, pyocyanin, and ROS. Each circle represents one TF, with the height indicating the significance, quantified by -log10 (p.adjust), and the size indicates the number of targets associated with the virulence pathway. (B). Intersection of master regulators in 6 virulence pathways. The bar chart on the top shows the number of intersections, while the matrix below indicates which TFs are involved in specific biological processes such as siderophore production, pyocyanin production, biofilm formation, ROS response, QS, and motility.

The intersection of master regulators across these virulent pathways revealed that six TFs (PA0167, PyrR, PA0707, PA0877, PA1504, and PA1484) only played an important role in regulating the motility pathway, and four TFs (PA1380, PA0815, PA5428, and PA3973) jointly regulated the motility and other pathways, including biofilm production, pyocyanin production, and QS pathways (Fig. 4B). Specifically, PA0815 co-regulated motility, biofilm production, and QS pathways. PA1380 jointly regulated motility and pyocyanin production pathways. PA3973 and PA5428 jointly regulated motility and biofilm production pathways. The regulatory network of these four master regulators and their target genes is summarised in Fig. S5, represented by lines connecting the TFs to the targets involved in specific pathways.

Earlier version: Predicted functions of regulators characterised in two metabolomic pathways

Revised version: Predicted functions of regulators characterised in two metabolomi pathways

Earlier version: P. aeruginosa uses different global regulators associated with metabolism that ensure its survival and enhance its adaptability under fluctuating environmental conditions. For instance, it uses different regulatory mechanisms to respond to nutrient changes, oxygen limitation, or CF-associated lung infection conditions45–48.

Revised version: P. aeruginosa uses different global regulators associated with metabolism that ensure its survival and enhance its adaptability under fluctuating environmental conditions. For instance, it uses different regulatory mechanisms to respond to nutrient changes, oxygen limitation, or CF-associated lung infection conditions50–53.

To define the core TFs that regulate metabolism, we performed master regulator analysis in metabolic pathways (p.adjust < 0.05) based on the TFBSs located in the promoter regions revealed by ChIP-seq for all TFs. We identified 14 TFs potentially involved in two pathways, namely the tricarboxylic acid (TCA) cycle and ribosome function pathways (Fig. S6A-B). The radial plot highlights the master regulators within the TCA cycle and ribosome function pathways, and their joint regulation with their targets is presented in the regulatory network below. PA2957 and PA5438 were found to play putative roles in the TCA cycle by potentially binding to the promoter of certain genes, including icd, sucA, and sdhC, indicating that they coordinate the expression of genes crucial for central metabolism and energy production (Fig. S6B). Furthermore, certain TFs, such as PA0611, PA0403 (PyrR), and PA1283, were found to be involved in regulating a broad array of ribosomal genes, such as rpsB, rpsA, and rplL. Notably, TF PyrR was found to play a putative role in regulating motility in P. aeruginosa. Overall, these regulators, which are related to the core metabolic pathways that enhance the adaptation of P. aeruginosa to different infection conditions and contribute to its pathogenicity, could offer potential drug targets for tackling P. aeruginosa infection in the future.

A global summary of transcription regulatory networks and functions in P. aeruginosa

After studying the TFBSs of most of the TFs in P. aeruginosa both in vivo and in vitro, we organised these datasets to yield a comprehensive connection of virulence-related master regulators. The QS and motility pathways consisting of 18 and 20 TFs, respectively, were the key pathways in the virulent landscape and showed a close collaboration with the other five pathways (Fig. 5A). Four TFs (ExsA, CprR, AlpR, and PA2449) were involved in T3SS, a direct host infection mechanism used by P. aeruginosa. Additionally, six TFs (PA4436, PA0929, PA3771, PA2534, PA3699, and PA5218) were involved in siderophore production, as identified using HT-SELEX, and showed correlations with only one TF, PA2534, involved in the ROS resistance pathway. Four TFs (CprR, PA0708, PA1520, and PA2479) were characterised in the stringent response and persister cell pathways and were correlated with other virulence-related pathways, namely QS, ROS resistance, T3SS, and antibiotic resistance pathways.

Overview of transcriptional network of TFs in P. aeruginosa.

(A). The interaction network of virulence-related master regulators in P. aeruginosa. The ten virulence pathways are highlighted in different colors. The size of the circle represents the degree, and the width of the edges indicates the overlapped number of TFs between two pathways. (B). The target annotated using the COG database shows four orthology classification profiling from different DBD types TFs. The size of rectangle indicates the number of targets.

Earlier version: To investigate the phylogenetic classification of the binding targets of the identified TFs on a genome-wide scale, we derived a global atlas revealing the relationships between the protein functions of target genes and the family categories of TFs (Fig. 5B). The Clusters of Orthologous Groups of proteins (COGs) database was used to classify proteins based on the orthology concept49. Based on the annotated COGs of P. aeruginosa, we found that most of the TF targets were involved in metabolism, followed by poorly characterised categories. The P. aeruginosa PAO1 genome is 6.3 Mbp long, and approximately 40% of it remains uncharacterised but is believed to encode well-conserved hypothetical proteins (HPs), which might have indispensable and similar functions. Given that most of the TF targets are poorly characterised, it is crucial to study these HPs in the future to provide more insights into P. aeruginosa pathogenicity. Altogether, our results provide a resource of TFs and their putative target genes in P. aeruginosa and, consequently, a basis for understanding the regulatory network of virulence and orthological functions.

Revised version: To investigate the phylogenetic classification of the binding targets of the identified TFs on a genome-wide scale, we derived a global atlas revealing the relationships between the protein functions of target genes and the family categories of TFs (Fig. 5B). The Clusters of Orthologous Groups of proteins (COGs) database was used to classify proteins based on the orthology concept54. Based on the annotated COGs of P. aeruginosa, we found that most of the TF targets were involved in metabolism, followed by poorly characterised categories. The P. aeruginosa PAO1 genome is 6.3 Mbp long, and approximately 40% of it remains uncharacterised but is believed to encode well-conserved hypothetical proteins (HPs), which might have indispensable and similar functions. Given that most of the TF targets are poorly characterised, it is crucial to study these HPs in the future to provide more insights into P. aeruginosa pathogenicity. Altogether, our results provide a resource of TFs and their putative target genes in P. aeruginosa and, consequently, a basis for understanding the regulatory network of virulence and orthological functions.

A Web-based database of all TFBSs in P. aeruginosa

Earlier version: To enable a quick search of the TFBSs in P. aeruginosa, we developed a Web-based database containing our ChIP-seq and HT-SELEX results at https://jiadhuang0417.shinyapps.io/PATF_Net/, which is available to the public. It has four parts: homepage, data page, upload, and contact. The homepage briefly describes the background information and schematic workflow of ChIP-seq and HT-SELEX. The data page provides two types of data: a network plot of TFs and targets, and detailed TF–target information in a table format. The network plot illustrates the top 50 peaks for ease of visualisation. The code was generated using R, and the figures were created using BioRender (https://www.biorender.com/).

Revised version: To enable a quick search of the TFBSs in P. aeruginosa, we developed a Web-based database containing our ChIP-seq and HT-SELEX results at https://jiadhuang0417.shinyapps.io/PATF_Net/, which is available to the public. It has four parts: homepage, data page, upload, and contact. The homepage briefly describes the background information and schematic workflow of ChIP-seq and HT-SELEX. The data page provides two types of data: a network plot of TFs and targets, and detailed TF–target information in a table format. The network plot illustrates the top 50 peaks for ease of visualisation. The code was generated using R, and the figures were created using BioRender (ChIP-seq: Created in BioRender; HT-SELEX: Created in BioRender).

Overall, this platform provides several services, including individual TF–target network visualisation, detailed peak annotations in a table format, TF–target interaction visualisation, data download, and data upload for those who want to share their results on this online platform. The primary purpose of PATF_Net is to store, search, and mine valuable information on P. aeruginosa TFs for researchers investigating P. aeruginosa infection in the future.

Conservation and evolution of TFs in P. aeruginosa

Earlier version: Differences between species are mainly due to differences in genes, whereas essential genes are usually well conserved, especially in subspecies. To investigate the conservation of TFs in P. aeruginosa, we performed a pan-genome analysis of 100 strains containing model strains and clinical strains using the Roary software50. A total of 21,432 genes were identified across the 100 P. aeruginosa strains, of which only 4,544 genes were classified as core genes (existing in more than 99% of strains), revealing that most of the genes were cloud genes (existing in less than 15% of strains) (Fig. 6A). However, when focusing on the TFs among these 100 strains, almost 89% of TFs (331 of 373) were found to be well-conserved (Fig. 6B), which implied that these recognised TFs from PAO1 play essential roles in biological activities in all of the P. aeruginosa subspecies.

Revised version: Differences between species are mainly due to differences in genes, whereas essential genes are usually well conserved, especially in subspecies. To investigate the conservation of TFs in P. aeruginosa, we performed a pan-genome analysis of 100 strains containing model strains and clinical strains using the Roary software55. A total of 21,432 genes were identified across the 100 P. aeruginosa strains, of which only 4,544 genes were classified as core genes (existing in more than 99% of strains), revealing that most of the genes were cloud genes (existing in less than 15% of strains) (Fig. 6A). However, when focusing on the TFs among these 100 strains, almost 89% of TFs (331 of 373) were found to be well-conserved (Fig. 6B), which implied that these recognised TFs from PAO1 play essential roles in biological activities in all of the P. aeruginosa subspecies.

Conservation and variability of TFs in PAO1.

Earlier version: (A). The pie chart shows the proportions for all genes. (B). The pie chart of proportion for all TFs identified from PAO1. (C). The bar plot of proportion for non-core TFs. Core genes indicate genes existed in 99% ∼ 100% strains; Soft core genes indicate genes existed in 95% ∼ 99% strains; Shell genes indicate genes existed in 15% ∼ 95% strains; Cloud genes indicate genes existed in 0% ∼ 15% strains. (D). The conservation and evolutionary relationship of all 373 TFs in PAO1 among bacteria, archaebacteria, fungi, plants, and animals. The conservation value was normalized after blastp alignment82. The phylogenetic trees were constructed using MEGA1183 and plotted via R package ggtree84.

Revised version: (A). The pie chart shows the proportions of genes categorized by their presence across P. aeruginosa strains for all genes. (B). The pie chart shows the distribution of TFs identified from PAO1 across different conservation categories. (C). The bar plot of the proportion for non- core TFs. Genes are categorized based on their presence frequency across P. aeruginosa strains: Core genes (present in 99% ∼ 100% strains), Soft core genes (present in 95% ∼ 99% strains), Shell genes (present in 15% ∼ 95% strains), and Cloud genes (present in 0% ∼ 15% strains). (D). The conservation and evolutionary relationship of all 373 TFs in PAO1 among bacteria, archaebacteria, fungi, plants, and animals. The conservation value was normalized after blastp alignment86. The phylogenetic trees were constructed using MEGA1187 and plotted via R package ggtree88.

Earlier version: The bar plot provides detailed information on the rest of the non-core TFs in PAO1, in addition to several important TFs, such as LasR and ExsA, that are lacking in some strains (Fig. 7C). As is known, the lasR mutant of P. aeruginosa is associated with CF lung disease progression and is considered a marker of an early CF-adaptive phenotype51. ExsA is considered an important activator of the T3SS pathway in P. aeruginosa, which is characterised by an acute infection phenotype. Additionally, VqsM might have a relatively special function in PAO1 among P. aeruginosa subspecies in addition to playing an important role in the QS and antibiotic resistance pathways. Taken together, the core- and non-core TFs were both vital and worth studying for understanding the metabolism and virulence-related regulatory mechanisms in P. aeruginosa.

Revised version: The bar plot provides detailed information on the rest of the non-core TFs in PAO1, in addition to several important TFs, such as LasR and ExsA, that are lacking in some strains (Fig. 7C). As is known, the lasR mutant of P. aeruginosa is associated with CF lung disease progression and is considered a marker of an early CF-adaptive phenotype56. ExsA is considered an important activator of the T3SS pathway in P. aeruginosa, which is characterised by an acute infection phenotype. Additionally, VqsM might have a relatively special function in PAO1 among P. aeruginosa subspecies in addition to playing an important role in the QS and antibiotic resistance pathways. Taken together, the core- and non-core TFs were both vital and worth studying for understanding the metabolism and virulence-related regulatory mechanisms in P. aeruginosa.

Earlier version: To further investigate the conservation of TFs in P. aeruginosa, we re-analysed the ChIP-seq data of PhoB and RpoN in the PA14 strain52,53 and compared them with those in the PAO1 strain. Fig. S7A shows genome-wide binding peaks of PhoB and RpoN throughout the genomes of PAO1 and PA14. We found that in PA14 and PAO1, PhoB had 725 and 63 peaks, while RpoN yielded 363 and 1,238 peaks, respectively. Next, we used MEME54 to compare the motifs of PhoB and RpoN in PAO1 and PA14 based on their binding sequences from ChIP-seq (Fig. S7B). PhoB and RpoN shared a well-conserved motif in the two strains even though they showed different peak distributions in the genomes. They had different numbers of peaks, and of the 690 and 345 respective annotation targets of PhoB and RpoN in PA14, 114 and 66 targets, respectively, were unique to PA14. The rest of the targets were homologous between PAO1 and PA14. We next performed target annotation and observed that the genes regulated by these two TFs were significantly (Fisher’s test) overrepresented in the two tested strains (Fig. S7C). For example, RpoN has been found to play an important role in regulating T6SS in both PAO1 and PA1453,55. Overall, the conservation of the representative TFs PhoB and RpoN in PAO1 and PA14 indicates that conserved TFs may share similar functions in different species. This knowledge is expected to guide drug development for different infectious strains.

Revised version: To further investigate the conservation of TFs in P. aeruginosa, we re-analysed the ChIP-seq data of PhoB and RpoN in the PA14 strain57,58 and compared them with those in the PAO1 strain. Fig. S7A shows genome-wide binding peaks of PhoB and RpoN throughout the genomes of PAO1 and PA14. We found that in PA14 and PAO1, PhoB had 725 and 63 peaks, while RpoN yielded 363 and 1,238 peaks, respectively. Next, we used MEME59 to compare the motifs of PhoB and RpoN in PAO1 and PA14 based on their binding sequences from ChIP-seq (Fig. S7B). PhoB and RpoN shared a well-conserved motif in the two strains even though they showed different peak distributions in the genomes. They had different numbers of peaks, and of the 690 and 345 respective annotation targets of PhoB and RpoN in PA14, 114 and 66 targets, respectively, were unique to PA14. The rest of the targets were homologous between PAO1 and PA14. We next performed target annotation and observed that the genes regulated by these two TFs were significantly (Fisher’s test) overrepresented in the two tested strains (Fig. S7C). For example, RpoN has been found to play an important role in regulating T6SS in both PAO1 and PA1458,60. Overall, the conservation of the representative TFs PhoB and RpoN in PAO1 and PA14 indicates that conserved TFs may share similar functions in different species. This knowledge is expected to guide drug development for different infectious strains.

Earlier version: TFs play a pivotal role in regulating gene expression, and studying their conservation and evolution could provide insights into how gene expression patterns are conserved or modified across different species and help to predict gene regulatory networks in different species. For example, the conservation of a TF across multiple species suggests that its targets and regulatory interactions are likely to be conserved as well. Therefore, to study the conservation and evolution of P. aeruginosa TFs across other species, we conducted phylogenetic analyses of the 373 TFs of P. aeruginosa across other bacteria, fungi, archaea, plants, and animals. We found that three P. aeruginosa species showed the most conserved TFs, followed by P. savastanoi, a phytopathogen (Fig. 6D). Compared with bacteria, the conservation values of the TFs were lower in fungi, archaea, plants, and animals. However, several orthologues of TFs were still noted among these organisms. For instance, PA2032 had more than 25% and approximately 30% identity with aminoadipate aminotransferase in humans and mice, respectively. Altogether, the high-level evolutionary conservation of TFs in P. aeruginosa strongly suggests that the regulatory mechanisms are common to a wide range of bacteria.

Revised version: TFs play a pivotal role in regulating gene expression, and studying their conservation and evolution could provide insights into how gene expression patterns are conserved or modified across different species and help to predict gene regulatory networks in different species. For example, the conservation of a TF across multiple species suggests that its targets and regulatory interactions are likely to be conserved as well. Therefore, to study the conservation and evolution of P. aeruginosa TFs across other species, we conducted phylogenetic analyses of the 373 TFs of P. aeruginosa across other bacteria, fungi, archaea, plants, and animals. We found that three P. aeruginosa species showed the most conserved TFs, followed by P. savastanoi, a phytopathogen (Fig. 6D). Compared with bacteria, the conservation values of the TFs were lower in fungi, archaea, plants, and animals. However, several orthologues of TFs were still noted among these organisms. For instance, PA2032 had more than 25% and approximately 30% identity with aminoadipate aminotransferase in humans and mice, respectively. The phylogenetic tree of PA2032 across bacteria, archaea, fungi, plants, and animals, with PAO1 as the root revealed that the bacterial TFs (purple) indicates a high degree of conservation within prokaryotes, suggesting a fundamental role in core regulatory processes (Fig. S7D). In contrast, eukaryotic TFs (fungi, plants, and animals) form distinct clades with longer branch lengths, indicating significant divergence and specialization during eukaryotic evolution. These findings suggest that while TF is conserved across domains of life, its functional roles and regulatory mechanisms have undergone substantial diversification in eukaryotes. Altogether, the high-level evolutionary conservation of TFs in P. aeruginosa strongly suggests that the regulatory mechanisms are common to a wide range of bacteria.

Discussion

Earlier version: Bacterial pathogens use many mechanisms to establish infection and cause diseases in human hosts, making them a major public health concern worldwide. These bacteria can secrete a wide range of molecular particles that recognise and bind to host cell targets, thus damaging or preventing host responses56. Pathogenic bacteria use strategies to adapt to environmental stressors such as nutrient limitation or antibiotic treatment57. The underlying mechanism of these strategies involves sensing external signals and then responding to different environmental conditions. TFs are key molecules involved in responding to host or environmental signals by precisely modulating transcription levels; hence, their pivotal roles in infection conditions cannot be overstated. Nevertheless, our comprehension of bacterial TFs remains considerably limited, even in extensively examined model organisms such as P. aeruginosa. To address this knowledge gap and shed more light on TFs in P. aeruginosa, we conducted ChIP-seq experiments aimed at elucidating the TFBSs for 172 TFs with relatively unexplored biological functions.

Revised version: Bacterial pathogens use many mechanisms to establish infection and cause diseases in human hosts, making them a major public health concern worldwide. These bacteria can secrete a wide range of molecular particles that recognise and bind to host cell targets, thus damaging or preventing host responses61. Pathogenic bacteria use strategies to adapt to environmental stressors such as nutrient limitation or antibiotic treatment62. The underlying mechanism of these strategies involves sensing external signals and then responding to different environmental conditions. TFs are key molecules involved in responding to host or environmental signals by precisely modulating transcription levels; hence, their pivotal roles in infection conditions cannot be overstated. Nevertheless, our comprehension of bacterial TFs remains considerably limited, even in extensively examined model organisms such as P. aeruginosa. To address this knowledge gap and shed more light on TFs in P. aeruginosa, we conducted ChIP-seq experiments aimed at elucidating the TFBSs for 172 TFs with relatively unexplored biological functions.

Earlier version: The basic function of TFs is to recognise specific DNA sequences (i.e., motifs) at the promoter region and then upregulate or downregulate the transcription of the target gene58. Apart from promoter regions, TFs have also been identified to directly regulate the transcription of coding regions in bacteria59. ChIP-seq is a valuable technique to obtain high-resolution DNA–protein interaction mapping on a genome- wide scale to study the potential binding targets of TFs. The technique is highly sensitive and can detect low-abundance DNA-binding events, allowing the detection of weak TF–DNA interactions. In vitro assays, such as DAP-seq or HT-SELEX, offer superior scalability to ChIP-seq but present their own challenges29,60. It is difficult to investigate the crosstalk and interactions between TFs by studying proteins individually. Thus, we devised a strategy involving TF-overexpressing plasmids. This strategy can enhance the visualisation of more subtle binding peaks, particularly when investigating co-associations and TFs lacking well-defined signals. Notably, the literature underscores the collaborative nature of TFs, which often exhibit interactions with other auxiliary proteins at multiple promoter sites61. As shown in Figures 3 and S3, we present a TF-binding motifs clustering tree and a TF co-association regulatory network based on ChIP-seq data and core clusters of TFs with high co-association scores. Overall, our above-described analysis has illustrated a regulatory atlas to study the biological functions of P. aeruginosa TFs in specific contexts.

Revised version: The basic function of TFs is to recognise specific DNA sequences (i.e., motifs) at the promoter region and then upregulate or downregulate the transcription of the target gene63. Apart from promoter regions, TFs have also been identified to directly regulate the transcription of coding regions in bacteria64. ChIP-seq is a valuable technique to obtain high-resolution DNA–protein interaction mapping on a genome- wide scale to study the potential binding targets of TFs. The technique is highly sensitive and can detect low-abundance DNA-binding events, allowing the detection of weak TF–DNA interactions. In vitro assays, such as DAP-seq or HT-SELEX, offer superior scalability to ChIP-seq but present their own challenges29,65. It is difficult to investigate the crosstalk and interactions between TFs by studying proteins individually. Thus, we devised a strategy involving TF-overexpressing plasmids. This strategy can enhance the visualisation of more subtle binding peaks, particularly when investigating co-associations and TFs lacking well-defined signals. Notably, the literature underscores the collaborative nature of TFs, which often exhibit interactions with other auxiliary proteins at multiple promoter sites66. As shown in Figures 3 and S3, we present a TF-binding motifs clustering tree and a TF co-association regulatory network based on ChIP-seq data and core clusters of TFs with high co-association scores. Overall, our above-described analysis has illustrated a regulatory atlas to study the biological functions of P. aeruginosa TFs in specific contexts.

Earlier version: P. aeruginosa uses TFs to control biological functions, including metabolism and virulence. Given the public health burden of P. aeruginosa infection, many studies have investigated the roles of TFs in this pathogen. These studies have characterised the regulatory mechanisms of individual TFs, such as FleQ, AlgR, LasR, VqsM, and PvrA62–66. Additionally, a TRN comprising 690 genes and 1,020 regulatory interactions with six biological modules and main motifs was established in 2011 using published data67. Furthermore, we developed a PAGnet containing 20 key virulence-associated TFs based on ChIP-seq and RNA-seq data68. A previous study used machine learning to identify modules of a group of genes related to known TFs from published data69. Another study used DAP-seq to determine a regulatory network of 55 RRs in P. aeruginosa29. To further identify and characterise unknown TFs, we used HT-SELEX to describe the binding specificities of 182 TFs28. Compared with the abovementioned studies, the present work profiled a comprehensive regulatory network of P. aeruginosa, combining existing and new data with experimental verification.

Revised version: =P. aeruginosa uses TFs to control biological functions, including metabolism and virulence. Given the public health burden of P. aeruginosa infection, many studies have investigated the roles of TFs in this pathogen. These studies have characterised the regulatory mechanisms of individual TFs, such as FleQ, AlgR, LasR, VqsM, and PvrA40,67–70. Additionally, a TRN comprising 690 genes and 1,020 regulatory interactions with six biological modules and main motifs was established in 2011 using published data71. Furthermore, we developed a PAGnet containing 20 key virulence-associated TFs based on ChIP-seq and RNA-seq data72. A previous study used machine learning to identify modules of a group of genes related to known TFs from published data73. Another study used DAP-seq to determine a regulatory network of 55 RRs in P. aeruginosa29. To further identify and characterise unknown TFs, we used HT-SELEX to describe the binding specificities of 182 TFs28. Compared with the abovementioned studies, the present work profiled a comprehensive regulatory network of P. aeruginosa, combining existing and new data with experimental verification.

Earlier version: We confirmed the functions and provided potential molecular mechanisms of a number of previously identified TFs to explain their phenotypes. PA0797 is known to regulate the pqs system and pyocyanin production70. In the present study, it was also found to bind to the pqsH promoter region and its motif was visualised. PA5428 was found to bind to the promoter regions of aceA and glcB genes71, which was also demonstrated in our ChIP-seq results. PA4381 (CloR) was found to be associated with polymyxin resistance in a previous study72 and to be possibly related to ROS resistance in the present study. Furthermore, PA5032 plays a putative role in biofilm regulation and also forms an operon with PA5033, an HP associated with biofilm formation73.

Revised version: We confirmed the functions and provided potential molecular mechanisms of a number of previously identified TFs to explain their phenotypes. PA0797 is known to regulate the pqs system and pyocyanin production74. In the present study, it was also found to bind to the pqsH promoter region and its motif was visualised. PA5428 was found to bind to the promoter regions of aceA and glcB genes75, which was also demonstrated in our ChIP-seq results. PA4381 (CloR) was found to be associated with polymyxin resistance in a previous study76 and to be possibly related to ROS resistance in the present study. Furthermore, PA5032 plays a putative role in biofilm regulation and also forms an operon with PA5033, an HP associated with biofilm formation77.

Earlier version: Our study is the first to illustrate the hierarchical regulatory network and ternary motifs of and genome-wide co-association relationships among TFs in P. aeruginosa, contributing to a better understanding of the fundamental traits of pathogenicity of this species. Traditional bacterial eradication strategies involving antibiotics often focus on pathogen extermination, yet this approach can inadvertently foster selective growth advantages and precipitate the emergence of drug-resistant strains. Therefore, prioritising TF-targeted drugs could potentially alleviate selective growth pressure and enhance future preventative measures against P. aeruginosa infection. Overall, we gathered the binding site data of most of the uncharacterised TFs in P. aeruginosa and established a Web-based database for storing and searching this information. The TF regulatory mechanisms clarified in this database can help to decode the pathogenic mechanisms and guide the development of new drugs targeting virulence-related TFs to combat P. aeruginosa infection. The conservation and evolutionary patterns of TFs also suggest that these TFs perform similar functions in diverse species, thus underscoring the broader applicability of our research to studies investigating TFs in other microbes.

Revised version: Our study is the first to illustrate the hierarchical regulatory network and ternary motifs of and genome-wide co-association relationships among TFs in P. aeruginosa, contributing to a better understanding of the fundamental traits of pathogenicity of this species. Traditional bacterial eradication strategies involving antibiotics often focus on pathogen extermination, yet this approach can inadvertently foster selective growth advantages and precipitate the emergence of drug-resistant strains. Therefore, prioritising TF-targeted drugs could potentially alleviate selective growth pressure and enhance future preventative measures against P. aeruginosa infection. The extensive datasets generated in this study offer valuable insights into understanding and targeting P. aeruginosa pathogenicity. The genome-wide binding profiles can be systematically analyzed through our hierarchical regulatory network framework to decode complex virulence mechanisms. The virulence-related master regulators and core regulatory clusters identified in this study highlighted key nodes of transcriptional control. Understanding these regulatory relationships is particularly valuable for identifying targets whose modulation would significantly impact virulence while accounting for potential compensatory mechanisms. This knowledge base thus provides a foundation for developing targeted approaches to combat P. aeruginosa infections, moving beyond traditional antibiotic strategies toward more sophisticated interventions based on regulatory network manipulation.

Methods

Strains, plasmids, and primers

The bacterial strains, plasmids, and primers used in the present study are listed in Table S8. The P. aeruginosa PAO1 strain, and its derivatives were grown at 37 °C in LB (Luria-Bertani) broth with shaking at 220 rpm or on LB agar plates. Antibiotics were used for Escherichia coli at the following concentrations: kanamycin at 50 μg/ ml and carbenicillin at 60 μg/ml.

ChIP-seq and analysis

Earlier version: The chromatin immunoprecipitation (ChIP) procedures were modified from a previous study74. Briefly, for the VSVG-tagged, the ORF was amplified by PCR from the PAO1 genome and cloned into pAK1900 plasmid by HindIII/BamHI for the overexpressed TFs through HindIII site by using ClonEx- press MultiS One Step Cloning Kit (Vazyme, China). Wild-type P. aeruginosa containing empty pAK1900 or pAK1900-TF-VSV-G was cultured in LB medium supplemented at 37°C with shaking until mid-log phase (OD600 = 0.6), then we cross-linked the samples with 1% formaldehyde for 10 min. Subsequently, cross-linking was stopped by the addition of 125 mM glycine. Samples were centrifuged and washed thrice with a Tris buffer (20 mM Tris-HCl, pH 7.5, 150 mM NaCl). The resulting pellets were resuspended in 500 μl IP buffer (50 mM HEPES–KOH [pH 7.5], 150 mM NaCl, 1 mM EDTA, 1 % Triton X- 100, 0.1% sodium deoxycholate, 0.1% SDS, and mini-protease inhibitor cocktail (Roche), and then s The DNA was broken into pieces (100 to 300 bp) with an ultrasonic processor. Insoluble cellular debris was removed by centrifugation at 4 °C, and the supernatant was used as the input sample in IP experiments. We next added 25 μl agarose-conjugated anti-VSV antibodies (Sigma) in the IP buffer. Washing, reverse cross-linked, and purification of the ChIP DNA were conducted. We used the NEXTflex ChIP-Seq kit (Bio Scientific) to construct the DNA fragment library, and agarose gel was used to cut DNA fragments between 150 and 250 bp. After sequencing, the raw reads were trimmed by Trim Galore75 with default parameters. Then, the filtered reads were aligned to the PAO1 genome (NC_002516) using bowtie2. Only the unique reads after alignment will be used to perform peakcalling with MACS2 (P < 0.001). MEME-ChIP was used to identify consensus motifs with all peaks. TF target genes were then annotated by the R package ChIPpeakAnno.

Revised version: The chromatin immunoprecipitation (ChIP) procedures were modified from a previous study78. Briefly, for the VSVG-tagged, the ORF was amplified by PCR from the PAO1 genome and cloned into pAK1900 plasmid by HindIII/BamHI for the overexpressed TFs through HindIII site by using ClonEx- press MultiS One Step Cloning Kit (Vazyme, China). Wild-type P. aeruginosa containing empty pAK1900 or pAK1900-TF-VSV-G was cultured in LB medium supplemented at 37°C with shaking until mid-log phase (OD600 = 0.6), then we cross-linked the samples with 1% formaldehyde for 10 min. Subsequently, cross-linking was stopped by the addition of 125 mM glycine. Samples were centrifuged and washed thrice with a Tris buffer (20 mM Tris-HCl, pH 7.5, 150 mM NaCl). The resulting pellets were resuspended in 500 μl IP buffer (50 mM HEPES–KOH [pH 7.5], 150 mM NaCl, 1 mM EDTA, 1 % Triton X- 100, 0.1% sodium deoxycholate, 0.1% SDS, and mini-protease inhibitor cocktail (Roche), and then the DNA was broken into pieces (100 to 300 bp) with an ultrasonic processor. Insoluble cellular debris was removed by centrifugation at 4 °C, and the supernatant was used as the input sample in IP experiments. We next added 25 μl agarose-conjugated anti-VSV antibodies (Sigma) in the IP buffer. Washing, reverse cross-linked, and purification of the ChIP DNA were conducted. We used the NEXTflex ChIP-Seq kit (Bio Scientific) to construct the DNA fragment library, and agarose gel was used to cut DNA fragments between 150 and 250 bp. After sequencing, the raw reads were trimmed by Trim Galore79 with default parameters. Then, the filtered reads were aligned to the PAO1 genome (NC_002516) using bowtie2. Only the unique reads after alignment will be used to perform peakcalling with MACS2 (P < 0.001). MEME-ChIP was used to identify consensus motifs with all peaks. TF target genes were then annotated by the R package ChIPpeakAnno.

Identification of virulence-related master regulators

Earlier version: We first generated gene lists associated with nine pathways: QS, motility, biofilm production, antibiotic resistance, T6SS, T3SS, ROS resistance, pyocyanin, and siderophores. Then, we calculated the number of overlap genes between target genes of TF and genes involved in the above eight pathways. The statistical significance of the master regulator was identified using the Hypergeometric test (BH-adjusted P < 0.05). R package ggplot276 was used to visualize the result.

Revised version: We first generated gene lists associated with nine pathways80: QS, motility, biofilm production, antibiotic resistance, T6SS, T3SS, ROS resistance, pyocyanin, and siderophores. Then, we calculated the number of overlap genes between target genes of TF and genes involved in the above eight pathways. The statistical significance of the master regulator was identified using the Hypergeometric test (BH-adjusted P < 0.05). R package ggplot281 was used to visualize the result.

GO analysis

Earlier version: GO enrichment analysis of TFs target genes was identified using R package clusterProfile77. The enriched GO terms with BH-adjusted P < 0.05 were defined as significantly enriched.

Revised version: GO enrichment analysis of TFs target genes was identified using R package clusterProfile82. The enriched GO terms with BH-adjusted P < 0.05 were defined as significantly enriched.

Pan-genome analysis

Earlier version: The information on the model and clinical P. aeruginosa strains78,79 can be found in Table S9. Briefly, the genome sequence of all strains were annotated by prokka80 and the outputs were then used as input files for roary.

Revised version: The information on the model and clinical P. aeruginosa strains83,84 can be found in Table S9. Briefly, the genome sequence of all strains were annotated by prokka85 and the outputs were then used as input files for roary.

Statistical analysis

Student’s t-test was performed to analyze the RT-qPCR and lux production results using Microsoft Office Excel 2019. *P < 0.05, **P < 0.01, and ***P < 0.001, and results represent means ± SD. All experiments were repeated at least twice.

Removed in the revised version: Data availability

Removed in the revised version: The ChIP-seq data was uploaded to National Center for Biotechnology Information Gene Expression Omnibus database under accession GSE241603 and GSE271817. Analysis codes have been deposited at https://github.com/dengxinb2315/PS-PATRnet-code.